Ceci est un bloc note interractif. Vous pouvez taper du code dans les cellues qui commencent par `Entrée [ ]:` et exécuter ce code en tapant **`<Maj+Entree>`** ou en appuyant sur le bouton <button class="btn btn-default" title="Exécuter"><i class="fa-step-forward fa"></i><span class="toolbar-btn-label">Exécuter</span></button> dans la barre ci-dessus.

Vérifier bien que vous avez complété toutes les zones contenant `TAPEZ VOTRE CODE ICI` ou "VOTRE REPONSE", ainsi que votre nom et les éventuels membres du groupe ci-dessous:

```python
NOM = ""
COLLABORATEURS = ""
```

---

<center><h1>Algorithmes gloutons</h1></center>

## Généralités

Optimiser un problème, c'est déterminer les conditions dans lesquelles ce problème présente une caractéristique spécifique. Par exemple, déterminer le minimum ou le maximum d'une fonction est un problème d'optimisation. On peut également citer la répartition  optimale  de  tâches  suivant  des  critères  précis,  le  <a href="https://fr.wikipedia.org/wiki/Probl%C3%A8me_du_rendu_de_monnaie">problème  du  rendu  de  monnaie</a>,  le  <a href="https://fr.wikipedia.org/wiki/Probl%C3%A8me_du_sac_%C3%A0_dos">problème  du  sac  à  dos</a>,  la  recherche d'un  plus  court  chemin  dans  un  graphe,  le  <a href="https://fr.wikipedia.org/wiki/Probl%C3%A8me_du_voyageur_de_commerce">problème  du  voyageur  de  commerce</a>, ... 

De  nombreuses  techniques  informatiques  sont susceptibles d'apporter une solution exacte ou approchée à ces problèmes. Certaines de ces techniques, comme l'énumération exhaustive de toutes les solutions, ont un coût machine qui les rend souvent peu pertinentes au regard de contraintes extérieures imposées (temps de réponse de la solution imposé, moyens machines limités). Les techniques de <a href="https://fr.wikipedia.org/wiki/Programmation_dynamique">programmation dynamique</a> ou certains algorithmes numériques peuvent apporter une solution.

Les algorithmes gloutons constituent une  alternative  dont  le  résultat  n'est  pas  toujours  optimal.  Plus  précisément,  ces algorithmes déterminent une solution en effectuant successivement des choix locaux, jamais remis en cause. Au cours de la construction de la solution, l'algorithme résout une partie du problème puis se focalise ensuite sur le sous-problème restant à résoudre. Une différence essentielle avec la programmation dynamique est que celle-ci peut remettre en cause des solutions déjà établies. Au lieu de se focaliser sur un seul sous-problème, elle explore les solutions de tous les sous-problèmes pour les combiner finalement de manière optimale. 
Le principal avantage des algorithmes gloutons est leur facilité de mise en œuvre. En outre, dans certaines situations dites canoniques, il arrive qu'ils renvoient la solution optimale d'un problème.

## Rendu de monnaie

### Présentation du problème

Un  achat  dit  en  espèces  se  traduit  par  un  échange  de  pièces  et  de  billets.  Dans la  suite,  les  pièces  désignent  indifféremment les véritables pièces que les billets.

Supposons qu'un achat induise un rendu de 49 euros. Quelles pièces peuvent être rendues ? La réponse, bien qu'évidente, n'est pas unique. Quatre pièces de 10 euros, 1 pièce de 5 euros et deux pièces de 2 euros conviennent. Mais quarante-neuf pièces de 1 euros conviennent également ! Si la question est de rendre la monnaie avec un minimum de pièces, le problème change de nature. Mais la réponse reste simple : c'est la première solution proposée. Toutefois, comment parvient-on à un tel résultat ? Quels choix ont été faits qui optimisent le nombre de pièces rendus ? C'est le problème du rendu de monnaie dont la solution dépend du système de monnaie utilisé.

Dans le système monétaire français, considérons pour simplifier que les pièces ont pour valeurs 1, 2, 5, 10, 20, 50, 100 euros (nous nous intéressons seulement aux valeurs entières et oublions l'existence du billet de 500 euros). Rendre 49 euros avec un minimum de pièces est un problème d'optimisation. En pratique, sans s'en rendre compte généralement, tout individu met en œuvre un algorithme glouton. Il choisit d'abord autant de fois que possible la plus grandeur valeur de monnaie contenue dans 49 euros. En l'occurence, il choisit donc quatre fois une pièce de 10 euros. La somme de 9 euros restant à rendre, il choisit une pièce de 5 euros, puis deux pièces de 2 euros.

Cette stratégie gagnante pour la somme de 49 euros l'est-elle pour n'importe quelle somme à rendre ? On peut montrer que la réponse est positive pour le système monétaire français. Pour cette raison, un tel système de monnaie est qualifié de canonique. D'autres systèmes ne sont pas canoniques. L'algorithme glouton ne répond alors pas de manière optimale.

><h3 class='fa fa-cogs' style="color: darkorange"> A faire vous-même </h3>
>
>Imaginons un pays dans lequel les pièces ont pour valeur 1, 7 et 10. Quelle solution propose l'algorithme glouton pour une somme de 14 euros ?

<!-- grade_id: cell-9ecade9c277c49f0, points: 2, grade: true, solution: true -->
VOTRE REPONSE

>Quelle est la solution optimale (c'est à dire avec un minimum de pièces) ?

<!-- grade_id: cell-a0a3205d799c776d, points: 1, grade: true, solution: true -->
VOTRE REPONSE

La réponse à cette difficulté passe par la programmation dynamique, thème abordé en spécialité NSI de classe de terminale.

### Un algorithme glouton

Considérons un ensemble de $n$ pièces de monnaie de valeurs : $$v_1 < v_2 < \dots < v_n$$ avec $v_1=1$. On peut noter le système de pièces : $$S_n=\{v_1, v_2, \dots, , v_n\}$$

Désignons par $s$ la somme à rendre avec le minimum de pièces de $S_n$. L'algorithme glouton sélectionne la plus grande valeur $v_n$ de $S_n$ et la compare à $s$. 
* Si $v_n>s$, la pièce de valeur $v_n$ ne peut pas être utilisée. On reprend l'algorithme avec la même valeur de $s$ et le système de pièces $S_{n-1}$.
* Si $v_n \leqslant s$, la pièce $v_n$ peut être utilisée une première fois. Ce qui fait une première pièce à comptabiliser, de valeur $v_n$, la somme restant à rendre étant alors $s-v_n$. L'algorithme continue donc avec la même système de pièces $S_n$ et cette nouvelle somme à rendre $s-v_n$.

L'algorithme est ainsi répété tant que la somme à rendre n'est pas nulle.

**Remarque :** Il s'agit effectivement d'un algorithme glouton, la plus grande valeur de pièce étant systématiquement choisie si sa valeur est inférieure à la somme à rendre. Ce choix ne garantit en rien l'optimalité globale de la solution. Le choix fait est considéré comme pertinent et permet d'avancer plus avant dans le calcul. Toutefois, comme nous l'écrivions plus haut, si le système monétaire est canonique, la solution est optimale.

### Codage en Python


<!-- grade_id: cell-d188aae9f4e75519, points: 0, grade: false, solution: true -->
```python
def pieces_a_rendre(somme_a_rendre, systeme_monnaie):
    """Renvoie la liste des pièces choisies pour rendre la somme `somme_a_rendre`.
    `systeme_monnaie` est la liste des pièces du système de monnaie utilisé, triée par valeur croissante."""
    # liste des pièces à rendre (au début une liste vide)
    # TAPEZ VOTRE CODE ICI
    # indice de la première pièce (celle de plus grande valeur) à comparer à la somme à rendre
    # TAPEZ VOTRE CODE ICI
    # boucle suivant l'algorithme décrit ci-dessus
    # TAPEZ VOTRE CODE ICI
    return lst_pieces
```

<!-- grade_id: cell-a3fc728a2f5d7853, points: 5, grade: true, solution: false -->
```python
assert pieces_a_rendre(49, [1, 2, 5, 10, 20, 50, 100]) == [20, 20, 5, 2, 2]
assert pieces_a_rendre(0, [1, 2, 5, 10, 20, 50, 100]) == []
assert pieces_a_rendre(14, [1, 7, 10]) == [10, 1, 1, 1, 1]
```

## Un problème d'organisation

Des conférenciers sont invités à présenter leurs exposés dans une salle. Mais leurs disponibilités ne leur permettent d'intervenir qu'à  des  horaires  bien  définis. Le  problème  est de construire  un  planning  d'occupation  de  la  salle  avec le plus grand nombre possible de conférenciers.

Désignons par $n$, entier  naturel  non  nul,  le  nombre  de  conférenciers.  Chacun  d'eux,  identifié  par  une  lettre $C_i$,  où $i$ est  un  entier compris  entre  $1$  et $n$,  est  associé  à  un  intervalle  temporel $[d_i, f_i[$ où $d_i$ et $f_i$ désignent  respectivement  l'heure  de  début et  l'heure  de  fin  de  l'intervention.  Afin  de  dégager  une  tactique  de  résolution  du  problème,  commençons  par  analyser  plusieurs situations.

### Premières analyses

#### Situation 1

Quatre conférenciers peuvent intervenir aux intervalles temporels suivants :

$$ C_1 : [3;4[ \quad C_2 : [0;1[ \quad C_3[2;3[ \quad C_4 : [1,2[$$

Pour mieux visualiser la situation, on se donne la fonction suivante :

```python
import matplotlib.pyplot as plt
def disp_interventions(lst_interventions):
    """Affiche les intervalles correspondants aux interventions possibles des conférenciers.
    `lst_interventions` est une liste de triplets de la forme (deb, fin, nom) indiquant que le
        conférencier `nom` propose une intervention sur l'intervalle de temps [deb; fin[."""
    hauteur = 1
    fin_prec = -1
    for deb, fin, nom in lst_interventions:
        if deb < fin_prec:
            hauteur += 1
        fin_prec = fin
        inter, = plt.plot([deb, fin], [hauteur]*2, 'o-')
        plt.text((deb+fin)/2, hauteur+0.2, nom, color=inter.get_color(), ha='center')
    plt.ylim(0, hauteur+1)
    plt.yticks([])
    plt.grid()
```

La situation 1 peut alors être visualisée ainsi :

```python
lst_interventions1 = [[3, 4, '$C_1$'], [0, 1, '$C_2$'], [2, 3, '$C_3$'], [1, 2, '$C_4$']]
disp_interventions(lst_interventions1)
```

Une telle situation est simple puisque tous les conférenciers peuvent intervenir sur des créneaux horaires disjoints. Le planning est donc défini par la suite $[C_2, C_4, C_3, C_1]$ de conférenciers.

#### Situation 2

On considère encore quatre conférenciers pouvant intervenir aux intervalles temporels suivants :

$$ C_1 : [3;8[ \quad C_2 : [1;2[ \quad C_3[0;4[ \quad C_4 : [5,7[$$

><h3 class='fa fa-cogs' style="color: darkorange"> A faire vous-même </h3>
>
>Définir une liste `lst_interventions2` correspondant à cette situation, et afficher le graphique correspondant :

<!-- grade_id: cell-deef9fa081245879, points: 2, grade: true, solution: true -->
```python
# TAPEZ VOTRE CODE ICI
```

Cette fois ci, les intervalles ne sont pas disjoints (on dira que les intervalles, ou les conférenciers, ne sont pas *compatibles*) et tous les conférenciers ne pourront donc pas intervenir : il va falloir faire un choix. Plusieurs planning peuvent être définis. Par exemple :

$$[C_2;C_1] \quad \text{ou} \quad [C_2;C_4] \quad \text{ou} \quad [C_3;C_4] \quad \text{ou} \quad [C_3] \quad \text{ou} \quad \dots$$

Seules les trois premières solutions sont retenues puisque la dernière ne maximise pas le nombre de conférenciers choisis.

Pour faire ce choix, nous allons esayer de mettre en place une *stratégie gloutonne*. Celle-ci va faire un premier choix de conférencier en suivant un critère à préciser. Ce choix ne sera jamais remis en question et la même stratégie sera appliquée pour trouver les conférenciers suivants.

Pour arriver à ce résultat, il va falloir commencer par trier les conférenciers. Deux idées semblent possibles : trier par heure de début ou par heure de fin. Le premier conférencier choisi sera alors le premier de la liste. On pourra ensuite parcourir la liste dans l'ordre, en ajoutant à chaque fois un conférencier s'il est compatible avec le dernier ajouté.

Par exemple, le tri des conférenciers par ordre croissant d'heure de début donne $[C_3;C_2;C_1;C_4]$.
  * On commence par choisir $C_3$;
  * $C_2$ n'est pas compatible avec $C_3$ : on le passe;
  * $C_1$ n'est pas compatible avec $C_3$ : on le passe;
  * $C_4$ est compatible avec $C_3$ : on le prend.

Le planning obtenu est $[C_3;C_4]$.


><h3 class='fa fa-cogs' style="color: darkorange"> A faire vous-même </h3>
>
>Donner le planning obtenu en triant les conférenciers par heure de fin croissante.

<!-- grade_id: cell-2e81ab78aab2150c, points: 2, grade: true, solution: true -->
VOTRE REPONSE

#### Situation 3

On considère à présent trois conférenciers pouvant intervenir aux intervalles temporels suivants :

$$ C_1 : [0;3[ \quad C_2 : [1;2[ \quad C_3[2;3[$$

><h3 class='fa fa-cogs' style="color: darkorange"> A faire vous-même </h3>
>
>Définir une liste `lst_interventions3` correspondant à cette situation, et afficher le graphique correspondant :

<!-- grade_id: cell-8ae96fdf7cec9ed4, points: 1, grade: true, solution: true -->
```python
# TAPEZ VOTRE CODE ICI
```

>Donner les plannings obtenus en appliquant l'algorithme précédent avec chacun des tris possibles. Que peut-on en conclure concernant le tri par ordre d'heure de début ?

<!-- grade_id: cell-c15e3d0975e1f098, points: 3, grade: true, solution: true -->
VOTRE REPONSE

On pourrait démontrer (mais ça sort du cadre de ce cours) que l'algorithme obtenu en commençant par trier par ordre croissant d'heure de fin donne toujours une solution optimale. 

### Le tri par heure de fin croissante

Tout d'abord, il va falloir trier la liste des interventions par ordre croissant d'heure de fin. Or la fonction `sorted` appliquée à une liste de triplets `[deb, fin, nom]` va trier d'abord suivant les heures de début. On va donc utiliser un paramètre optionnel de cette fonction, le paramètre `key`, correspondant à une fonction à appliquer à chaque tuple de la liste avant le tri. Ici, on ne s'intéresse qu'au deuxième élément du tuple, donc on veut une fonction qui à une intervention `inter` (c'est à dire un triplet) associe l'heure de fin (c'est à dire son deuxième élément : `inter[1]`) :

```python
print("non trié : ", lst_interventions1)
print("    trié : ", sorted(lst_interventions1, key=lambda inter:inter[1]))
print("non trié : ", lst_interventions2)
print("    trié : ", sorted(lst_interventions2, key=lambda inter:inter[1]))
print("non trié : ", lst_interventions3)
print("    trié : ", sorted(lst_interventions3, key=lambda inter:inter[1]))
```

**Remarque :** Ce problème de tri est un peu compliqué. Si ça ne vous paraît pas entièrement clair, ce n'est pas très grave. Vous pouvez vous contenter de retenir que pour obtenir une liste `lst_triee` triée par heures de fin croissantes en partant d'une liste `lst`, on doit écrire `lst_triee = sorted(lst, key=lambda inter:inter[1])`

### Un exemple plus complet à la main

Considérons la situation ci dessous :

```python
lst_interventions4 = [[0,7,'C1'],[2,5,'C2'],[6,8,'C3'],[1,2,'C4'],[5,6,'C5'],[0,2,'C6'],[4,7,'C7'],[0,1,'C8'],[3,6,'C9'],[1,3,'C10'],[4,5,'C11'],[6,8,'C12'],[0,2,'C13'],[5,7,'C14'],[1,4,'C15']]
disp_interventions(lst_interventions4)
```

><h3 class='fa fa-cogs' style="color: darkorange"> A faire vous-même </h3>
>
>Définir une liste `lst_triee` correspondant au tri de cette liste par heure de fin croissante :

<!-- grade_id: cell-8a8fe42343422122, points: 0, grade: false, solution: true -->
```python
# TAPEZ VOTRE CODE ICI
```

<!-- grade_id: cell-c1738a4098b981ec, points: 1, grade: true, solution: false -->
```python
assert list(map(lambda inter:inter[2], lst_triee)) == ['C8', 'C4', 'C6', 'C13', 'C10', 'C15', 'C2', 'C11', 'C5', 'C9', 'C1', 'C7', 'C14', 'C3', 'C12']
disp_interventions(lst_triee)
```

><h3 class='fa fa-cogs' style="color: darkorange"> A faire vous-même </h3>
>
>Appliquer pas à pas l'algorithme pour déterminer un planning :

<!-- grade_id: cell-03f07337f1c3c653, points: 4, grade: true, solution: true -->
VOTRE REPONSE

### Codage en Python

><h3 class='fa fa-cogs' style="color: darkorange"> A faire vous-même </h3>
>
>Il ne reste plus qu'à écrire une fonction en Python correspondant à cet algorithme :

<!-- grade_id: cell-64d324c18057990c, points: 0, grade: false, solution: true -->
```python
def calc_planning(lst_interventions):
    """Renvoie un planning optimal, c'est à dire une liste d'interventions de conférenciers de longueur
    maximum correspondant à la liste d'interventions `lst_interventions` donnée."""
    # tri des intervalles par valeurs croissantes d'heures de fin
    # TAPEZ VOTRE CODE ICI
    # Initialisation de la liste `lst_planning` (vide au début)
    # TAPEZ VOTRE CODE ICI
    # heure de fin de la dernière intervention ajoutée (0 au début)
    # TAPEZ VOTRE CODE ICI
    # Parcours des interventions, en ajoutant celles compatibles avec la dernière ajoutée
    # TAPEZ VOTRE CODE ICI
    return lst_planning
```

<!-- grade_id: cell-04fc71063428846b, points: 5, grade: true, solution: false -->
```python
assert calc_planning(lst_interventions1) == [[0, 1, '$C_2$'], [1, 2, '$C_4$'], [2, 3, '$C_3$'], [3, 4, '$C_1$']]
assert calc_planning(lst_interventions2) == [[1, 2, '$C_2$'], [5, 7, '$C_4$']]
assert calc_planning(lst_interventions3) == [[1, 2, '$C_2$'], [2, 3, '$C_3$']]
assert calc_planning(lst_interventions4) == [[0, 1, 'C8'], [1, 2, 'C4'], [2, 5, 'C2'], [5, 6, 'C5'], [6, 8, 'C3']]
```

Vous pouvez tester votre fonction avec une liste d'interventions tirée aléatoirement :

```python
from random import randint
def inter_alea(nb_inter, debut, fin, duree_max):
    """Renvoie une liste de `nb_inter` interventions étalées entre les instants `debut` et `fin`,
    chaque intervention durant au maximum `duree_max`."""
    lst_inter = []
    for i_inter in range(nb_inter):
        duree = randint(1, duree_max)
        deb = randint(debut, fin - duree)
        nom = '$C_{' + str(i_inter + 1) +'}$'
        lst_inter.append([deb, deb + duree, nom])
    return lst_inter
```

```python
lst_interventions = inter_alea(25, 0, 15, 3)
disp_interventions(lst_interventions)
```

```python
planning = calc_planning(lst_interventions)
disp_interventions(planning)
```